# Find the best model

## Setup

```python
import pandas as pd

from sklearn.linear_model import LinearRegression
from sklearn.metrics import mean_squared_error
from sklearn.metrics import mean_absolute_error
```

## Data

### Create data

```python
df = pd.DataFrame(
    {'sales': [2500, 4500, 6500, 8500, 10500, 12500, 14500, 16500, 18500, 20500],
      'ads'  : [900, 1400, 3600, 3800, 6200, 5200, 6800, 8300, 9800, 10100]}
)
```

### Variable lists

We need to prepare our data for the scikit-learn model (which does not work with Pandas dataframes):

- We create the variable `y_label` (our dependent variable, also called label): choose "`sales`" (this step will be helpful in later notebooks)

- We select our features (independent variables) and save them as `X`: choose "`ads`"

- Select and save the dependent variable as `y`

Hint:

```pyhton
y_label = "___"

X = df[["___"]]
y = df[___]
```

<!-- grade_id: y_label, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
y_label = "sales"

X = df[["ads"]]
y = df[y_label]
### END SOLUTION 
```

```python
X
```

Output:
```
ads
0    900
1   1400
2   3600
3   3800
4   6200
5   5200
6   6800
7   8300
8   9800
9  10100
```

```python
y
```

Output:
```
0     2500
1     4500
2     6500
3     8500
4    10500
5    12500
6    14500
7    16500
8    18500
9    20500
Name: sales, dtype: int64
```

<!-- grade_id: c, points: 2, grade: true, solution: false -->
```python
### Check your code (note that there are some hidden checks)
assert len(X) == 10

### BEGIN HIDDEN TESTS 
assert y_label == "sales"
assert len(y) == 10
### END HIDDEN TESTS
```

## Model

### Select model

Hint:

Save the `LinearRegression()`estimator as object called `reg`

---

```pyhton
___ = ___()
```

---

<!-- grade_id: reg, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
reg = LinearRegression()
### END SOLUTION 
```

<!-- grade_id: c-reg, points: 2, grade: true, solution: false -->
```python
### Check your code
assert reg.get_params() == {'copy_X': True, 'fit_intercept': True, 'n_jobs': None, 'normalize': 'deprecated', 'positive': False}
```

### Fit model


Fit the model (your `reg` object) to the data with `.fit()`.

In the fit function, you first need to provide the features and then the dependent variable (features, dependent variable)

Hint:


---

```pyhton
___.___(___, ___)
```

---



<!-- grade_id: reg-fit, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
reg.fit(X, y)
### END SOLUTION
```

Output:
```
LinearRegression()
```

<!-- grade_id: c-reg-fit, points: 2, grade: true, solution: false -->
```python
### Check your code
assert reg.positive == False
```

### Coefficients

```python
# Intercept
reg.intercept_
```

Output:
```
array([1181.21008618])
```

```python
# Slope
reg.coef_
```

Output:
```
array([[1.83935649]])
```

### Make predictions


Use your model to:

- predict sales from your ads (`X`) values with the method `.predict()`

- Save the result as `y_pred`

Hint:

---

```pyhton
___ = ___.___(___)
```

---



<!-- grade_id: y-pred, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
y_pred = reg.predict(X)
### END SOLUTION
```

```python
y_pred[0]
```

Output:
```
2836.6309279665047
```

<!-- grade_id: c-y-pred, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 2836 < y_pred[0] <  2840
```

### Evaluation

#### Mean squared error

Use the scikit-learn function `mean_squared_error()` to calculate the mean squared error and save the result as `mse`

Hint:

---

```pyhton
___ = ___(___, ___)
```

---

<!-- grade_id: mse, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mse = mean_squared_error(y, y_pred)
### END SOLUTION
```

```python
print(f'The mean squared error equals {mse:.0f} ')
```

Output:
```
The mean squared error equals 1160739
```

<!-- grade_id: c-mse, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 1160738 < mse < 1160740 
```

#### Root mean squared error

Use the scikit-learn function `mean_squared_error()` to calculate the root mean squared error and save the result as `rmse`.

Hint:

---

```pyhton
___ = ___(___, ___, ___)
```

- use the option `squared=False` 

---

<!-- grade_id: rmse, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
rmse = mean_squared_error(y, y_pred, squared=False)
### END SOLUTION
```

```python
print(f'The root mean squared error equals {rmse:.0f} ')
```

Output:
```
The root mean squared error equals 1077
```

<!-- grade_id: c-rmse, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 1076 < rmse < 1078 
```

#### Mean absolute error

Use the scikit-learn function `mean_absolute_error()` to calculate the mean absolute error and save the result as `mae`.

Hint:

---

```pyhton
___ = ___(___, ___)
```


---

<!-- grade_id: mae, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mae = mean_absolute_error(y, y_pred)
### END SOLUTION
```

```python
print(f'The mean absolute error equals {mae:.0f} ')
```

Output:
```
The mean absolute error equals 886
```

<!-- grade_id: c-mae, points: 2, grade: true, solution: false -->
```python
### Check your code
assert 885 < mae < 887 
```